```python
# Initialize Otter
import otter
grader = otter.Notebook("assignment_07.ipynb")
```

# Exploring JupyterLab

In this assignment, we will explore the advanced features of JupyterLab - a web-based interactive development environment for working with Jupyter notebooks, code, and data. We will cover different cell types, markdown text formatting, opening additional views of the same notebook, rendering cell output on the side of a code cell, running shell commands, using magic commands, graphical debugger, and workspaces. These features can enhance your productivity when working on data science projects, and by the end of this assignment, you will have a better understanding of how to use them effectively.

<!-- BEGIN QUESTION -->

<div class="alert alert-info">

## Submission instructions

<p>You receive marks for submitting your lab correctly, please follow these instructions:</p>

<ul>
  <!-- <li>Make at least three commits.</li>-->
  <li>Export to HTML via <code>File -> Save and export notebook as... -> HTML</code> and upload the HTML file to Canvas.</li>
    <ul>
      <li>Before submitting, make sure you restart the kernel and rerun all cells.</li>
    </ul>
  <li>Don't change any variable names that are given to you, don't move cells around, and don't include any code to install packages in the notebook.</li>
  </li>
</ul>
</div>

_Points:_ 1

<!-- END QUESTION -->

----

## Recap

Let's practice again skills we have learnt in a previous Modules:

1.  Create a new GitHub repo called `jupyterlab_practice`

    - make repo `Public`
    - initialize a repo with `README` file
    - Choose a license : `MIT License`  

2.  After the repo is created, add `.gitignore` file and populate it with `**/*.log`, `.DS_Store` and `v`
 
 
3. Clone the `jupyterlab_practice` repo to your computer

    - create a branch called `developer-1` 
    - set up the tracking branch for your local `developer-1` branch 
    - **use `developer-1` branch to make any changes to your repo** 
    - publish the branch on the remote

4. Create an environment (you can use any name) and export it as `environment.yaml` file. Make sure that exported environment can be used on different operating systems


For this exercise you will be working with your `jupyterlab_practice environment.`


**Don't forget to do frequent commits to track your changes**

<div class="alert alert-info">
    
### Question 1
    
You've recently installed JupyterLab in your `base` environment, and created a few new Conda environments. Unfortunately, you can't find your new environments in the list of available kernels in JupyterLab's Launcher page. What could be the reason? Select all correct answers.
        
A) `nb_conda_kernels` is not installed in the new environments.   
B) `ipykernel` is not installed in your new environment.  
C) JupyterLab is not installed in the new environments.    
D) `ipykernel` is not installed in your base environment.  
E) You haven't activated the new environment in your terminal before launching JupyterLab.   
F) `nb_conda_kernels` is not installed in your base environment.

</td>
</tr>
</table>
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q1 = "BF"
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 2
    
The terminal that you can open in JupyterLab can only use the Conda environment of your notebook's kernel.

A) True     
B) False    
 

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q2 = "B"
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

-----

## Part I: Learning More about JupyterLab

Let's start with reviewing some theory :

<div class="alert alert-info">
    
### Question 3

In JupyterLab, you can **open** and **edit** the following types of files (select all correct answers):



A)  `.md` (Markdown documents)  
B)  `.pdf` (PDF files)  
C)  `.py` (Python files)  
D)  `.ipynb` (Jupyter notebooks)  
E)  `.png` (image files)
 
<details><summary><b>Hint:</b></summary>
    
To answer this question, try and open <code>.md</code>,<code>.pdf</code>,<code>.py</code>, <code>.ipynb</code> and <code>.png</code> files and check whether 1) you can open them in Jupyter Lab and 2) you can edit content inside the file

</details>
</div>
 

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q3 = "ACD"
```

```python
grader.check("q3")
```

Output:
```
q3 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 4 

Running a code cell in a Jupyter notebook may change the value of a variable. This automatically updates the output of previously executed cells according to the new value of the changed variable.

A) True     
B) False 

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q4 = "B"
```

```python
grader.check("q4")
```

Output:
```
q4 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 5

Suppose that you have selected a cell in a Jupyter notebook (you can use the cell above), and you press `D` on your keyboard **twice**. What do you think would happen?
    
<details><summary><b>Hint:</b></summary>
    
Make sure your cell is in correct mode: either command or edit. you can see the Mode at teh bottom of your Jupyter Lab screen.
    
</details>
<br>
    
A) If you're in command mode, "dd" will be typed in the cell.     
B) If you're in edit mode, the cell will get deleted.   
C) If you're in command mode, the cell will get deleted.     
D) The cell will get cut and will be ready to be pasted elsewhere.   
E)  If you're in edit mode, "dd" will be typed in the cell.   




</div>


_Points:_ 1

```python
# You can test the keyboard shortcut on this cell
# but make sure you don't delete the solution cell below
```

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q5 = "CE"
```

```python
grader.check("q5")
```

Output:
```
q5 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 6

When you run a Markdown cell in a Jupyter notebook, the formatted text will be rendered and any code will be executed.

<details><summary><b>Hint:</b></summary>
    
Try to add a markdown cell above and add some code inside a cell. Run a cell to see what happens with the output 
    
</details>
<br>
 
    
A)	TRUE   
B)  FALSE   
 
  
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q6 = "B"
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 7

You're working on a data science project and you need to keep an eye on two plots generated by your code cells and the log generated by one other cell at all times. Would you use notebook views or output views, and why?
    
A) Notebook views, because all views use the same kernel.   
B) Notebook views, because we can point them to the preferred location in the notebook.  
C) Output views, because we can re-run the corresponding cell at any time.      
D) Output views, because they stay fixed as the notebook gets updated with outputs from other cells.  
</div>

_Points:_ 1

```python
# Your answer should be a string with one letter
q7 = "D"
```

```python
grader.check("q7")
```

Output:
```
q7 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 8 

Below you can see five cells with Python code from a Jupyter Notebook. In which sequence would we need to execute these five cells in order for the fifth one to print the result shown below `(8 ,2)`?

<img src="images/pic01.png" width=50% >
    
<details><summary><b>Hint:</b></summary>
    
You can recreate these cells shown on the image and run the cells in an order specified. Then check the final output.
    
</details>

    
A) 1, 2, 3, 4, 1, 4, 5   
B) 1, 2, 3, 4, 5   
C) 1, 2, 3, 5   
D) 1, 2, 3, 4, 3, 4, 5  
E) 1, 2, 3, 4, 1, 3, 5   

</div>

_Points:_ 1

```python
a =0
b=1
x=2
for i in range(2):
    a = a+1
    x = x+a*b
a=0
for i in range(2):
    a = a+1
    x = x+a*b

print(f'{x}, {a}')

```

Output:
```
8, 2
```

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q8 = "A"
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```


<div class="alert alert-info">
    
### Question 9

`%%timeit` is a convenient way to estimate the average running time of a code cell written in Julia, Python, or R in a Jupyter notebook.

A)  TRUE   
B)  FALSE   

</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q9 = "A"
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

<div class="alert alert-info">
    
### Question 10

You are working in a Jupyter notebook and you want to quickly experiment with one of your Python functions to make sure that it produces correct outputs without cluttering your notebook. Which JupyterLab feature would you use for this purpose?

A)  Jupyter extensions       
B)  Notebook views        
C)  Magic commands      
D)  JupyterLab terminal      
E)  Notebook console    
    
</div>

_Points:_ 1

```python
# Your answer should be a string with all the letters of the options you think are correct, e.g. 'AB'
q10 = 'E'
```

```python
grader.check("q10")
```

Output:
```
q10 results: All test cases passed!
```

----

## Part II: Create a new Jupyter Notebook 

<!-- BEGIN QUESTION -->


<div class="alert alert-info">
    
### Question 11


Create a Jupyter Notebook called `question11.ipynb` that contains the following cells:

1. A markdown cell with a level 1 heading that contains your name and one sentence under the heading saying "Notebook for assignment 7".
2. A code cell with a Python function that computes the mean of the values in a list and prints the results. The function should be called at the bottom of the code cell with the syntax `compute_mean([2, 4, 6, 8])` which should output `5` (or `5.0`). The code cell should be executed so that the output is shown.
3. Create a markdown cell under the code cell that explains what your code does and includes the equation of the mean written in Latex. To practice your Latex, please include two different ways of expression the equation of the mean ($\bar{x}$) as shown below:
    
<img src="images/pic02.png" width=50% >
    
This might look intimidating at first, but as we mentioned during lectures, it is common to not know the exact command for each part of an equation that you want to include and look it up online. A good place to start can be these references:

- https://en.wikipedia.org/wiki/Summation#Capital-sigma_notation (explains the sum notation)
- https://davidhamann.de/2017/06/12/latex-cheat-sheet (some common latex formulas)
- https://latex-tutorial.com/big-parentheses (how to make large parentheses)
    
    
Once you are done, take a single screenshot of all three cells from your Jupyter Notebook. Upload the screenshot to github and add a link to that screenshot in the cell below. 

    
</div>

_Points:_ 2

https://github.com/yuyangsu2010/jupyterlab_practice/blob/developer-1/function_screenshot.png

<!-- END QUESTION -->

Once you done, click the â–¶â–¶ button or go to `Kernel -> Restart Kernel and Run All Cells...` in the menu. Then stage and commit the changes. Push them to the remote

<!-- BEGIN QUESTION -->


<div class="alert alert-info">
    
### Question 12

Create a Jupyter notebook called `question12.ipynb` containing the following cells:

1. A markdown cell containing the text `We are currently working in the following directory:`
2. A code cell containing the shell command for displaying the current working directly using the special syntax that allows us to access shell commands in Jupyter Lab cells. This cell should be executed so that the current directory shows up as output.
3. Another markdown cell with the text: <code>To time how long it takes to run a Python cell, we could use the` `timeit` `command inside a cell as is shown below:<code> Followed by the proper use of the `timeit` command to time the following loop:
    
<details><summary><b>Note:</b></summary>
The code below should <b>not</b> be written as a separate code cell. Instead it should be included inside the markdown cell as a code block. Use Python syntax highlighting to make sure that the code is easy to read.
</details>
    
    ```
    for i in range(1_000_000):
        i**2
    ```


4. A code cell that contains the command to export the Jupyter notebook to HTML without including any of the input code cells.
5. Run the entire notebook to create the HTML exported version. Open the HTML exported notebook in a new JupyterLab tab and position this tab below your Jupyter notebook so that the view is split in half horizontally with the ipynb file in the top half and the HTML file in the bottom half.

**Once you are done, take a screenshot showing that you have the correct content in the notebook and that you have correctly positioned both tabs inside JupyterLab. Upload the screenshot to github and add a link in the cell below.**

</div>

_Points:_ 2

https://github.com/yuyangsu2010/jupyterlab_practice/blob/developer-1/screeshot_question12_and_html_export.png


<!-- END QUESTION -->

Once you done, click the â–¶â–¶ button or go to `Kernel -> Restart Kernel and Run All Cells...` in the menu. Then stage and commit the changes. Push them to the remote

<!-- BEGIN QUESTION -->

<div class="alert alert-info">
    
### Question 13

Once you done, commit and push your changes to the remote. Create a pull request and merge your branch into `main` branch. Send us a URL link to your `jupyterlab_practice` repo. 

You will be evaluated based on:

- presence of all required files
- comprehensive README.md
- .gitignore file
- quality of commit message
- completed pull request

</div>

_Points:_ 3

_Type your answer here, replacing this text._

<!-- END QUESTION -->

<div class="alert alert-danger" style="color:black">
    
**Restart and run all before submitting**
    
Before submitting,
don't forget to run all cells in your notebook
to make sure there are no errors
and so that the TAs can see the output of all the cells properly.
You can do this by clicking the â–¶â–¶ button
or going to `Kernel -> Restart Kernel and Run All Cells...` in the menu.
This is not only important for this course,
but a good habit you should get into before ever committing a notebook to GitHub,
so that your collaborators can run it from top to bottom
without issues.
</div>

------

Assignment 7 completed, Congratulations! â­

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Output:
```
q1 results: All test cases passed!

q10 results: All test cases passed!

q2 results: All test cases passed!

q3 results: All test cases passed!

q4 results: All test cases passed!

q5 results: All test cases passed!

q6 results: All test cases passed!

q7 results: All test cases passed!

q8 results: All test cases passed!

q9 results: All test cases passed!
```

Tests for q1 (points: 1)
```
>>> assert 'f' in q1.lower()
>>> assert 'b' in q1.lower()
>>>
```
failure_message: You are missing one of the correct answers
```
>>> assert 'a' not in q1.lower()
>>> assert 'c' not in q1.lower()
>>> assert 'd' not in q1.lower()
>>> assert 'e' not in q1.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q10 (points: 1)
```
>>> assert 'e' == q10.lower()
>>>
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q2 (points: 1)
```
>>> assert 'b' == q2.lower()
```
failure_message: You chose an incorrect answer


Tests for q3 (points: 1)
```
>>> assert 'a' in q3.lower()
>>> assert 'd' in q3.lower()
>>> assert 'c' in q3.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'b' not in q3.lower()
>>> assert 'e' not in q3.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q4 (points: 1)
```
>>> assert 'b' == q4.lower()
>>>
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q5 (points: 1)
```
>>> assert 'c' in q5.lower()
>>> assert 'e' in q5.lower()
```
failure_message: You are missing one of the correct answers
```
>>> assert 'a' not in q5.lower()
>>> assert 'b' not in q5.lower()
>>> assert 'd' not in q5.lower()
```
failure_message: At least one of the options you have selected is not correct


Tests for q6 (points: 1)
```
>>> assert 'b' == q6.lower()
```
failure_message: You are missing one of the correct answers


Tests for q7 (points: 1)
```
>>> assert q7.upper() == 'D'
```
failure_message: Your answer is incorrect or you are missing one of the correct answers


Tests for q8 (points: 1)
```
>>> assert 'a' == q8.lower()
```
failure_message: You chose an incorrect answer


Tests for q9 (points: 1)
```
>>> assert 'a' == q9.lower()
```
failure_message: Your answer is incorrect or you are missing one of the correct answers
